# Assignment 3

In this assignment you will explore measures of centrality on two networks, a friendship network in Part 1, and a blog network in Part 2.

## Part 1

Answer questions 1-4 using the network `G1`, a network of friendships at a university department. Each node corresponds to a person, and an edge indicates friendship. 

*The network has been loaded as networkx graph object `G1`.*

```python
import networkx as nx

G1 = nx.read_gml('assets/friendships.gml')
```

### Question 1

Find the degree centrality, closeness centrality, and betweeness centrality of node 100.

*This function should return a tuple of floats `(degree_centrality, closeness_centrality, betweenness_centrality)`.*

<!-- grade_id: cell-bffef7d5151cb470, points: 0, grade: false, solution: true -->
```python
def answer_one():
    
    degree_c = nx.degree_centrality(G1)[100]
    closeness_c = nx.closeness_centrality(G1)[100]
    betweenness_c = nx.betweenness_centrality(G1)[100]

    return (degree_c, closeness_c, betweenness_c)

answer_one()
```

Output:
```
(0.0026501766784452294, 0.2654784240150094, 7.142902633244772e-05)
```

<!-- grade_id: cell-cd6a99ae1fdd71a9, points: 1, grade: true, solution: false -->
```python
ans_one = answer_one()
assert type(ans_one) == tuple, "You must return a tuple"

```

### Use centrality measures to answer questions 2-4

### Question 2

Suppose you are employed by an online shopping website and are tasked with selecting one user in network G1 to send an online shopping voucher to. We expect that the user who receives the voucher will send it to their friends in the network.  You want the voucher to reach as many nodes as possible. The voucher can be forwarded to multiple users at the same time, but the travel distance of the voucher is limited to one step, which means if the voucher travels more than one step in this network, it is no longer valid. Apply your knowledge in network centrality to select the best candidate for the voucher. 

*This function should return an integer, the chosen node.*

```python
nx.closeness_centrality?
```

Output:
```
[0;31mSignature:[0m [0mnx[0m[0;34m.[0m[0mcloseness_centrality[0m[0;34m([0m[0mG[0m[0;34m,[0m [0mu[0m[0;34m=[0m[0;32mNone[0m[0;34m,[0m [0mdistance[0m[0;34m=[0m[0;32mNone[0m[0;34m,[0m [0mwf_improved[0m[0;34m=[0m[0;32mTrue[0m[0;34m)[0m[0;34m[0m[0;34m[0m[0m
[0;31mDocstring:[0m
Compute closeness centrality for nodes.

Closeness centrality [1]_ of a node `u` is the reciprocal of the
average shortest path distance to `u` over all `n-1` reachable nodes.

.. math::

    C(u) = \frac{n - 1}{\sum_{v=1}^{n-1} d(v, u)},

where `d(v, u)` is the shortest-path distance between `v` and `u`,
and `n-1` is the number of nodes reachable from `u`. Notice that the
closeness distance function computes the incoming distance to `u`
for directed graphs. To use outward distance, act on `G.reverse()`.

Notice that higher values of closeness indicate higher centrality.

Wasserman and Faust propose an improved formula for graphs with
more than one connected component. The 
```

<!-- grade_id: cell-f6296c30200cc9d0, points: 0, grade: false, solution: true -->
```python
def answer_two():
    
    degree_c = nx.degree_centrality(G1)
    max_degree_c = max(degree_c, key=lambda x: degree_c[x])
    
    return max_degree_c

answer_two()
```

Output:
```
105
```

<!-- grade_id: cell-d890b05007b8cce5, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()

```

### Question 3

Now the limit of the voucher’s travel distance has been removed. Because the network is connected, regardless of who you pick, every node in the network will eventually receive the voucher. However, we now want to ensure that the voucher reaches nodes as quickly as possible (i.e. in the fewest number of hops). How will you change your selection strategy? Write a function to tell us who is the best candidate in the network under this condition.

*This function should return an integer, the chosen node.*

<!-- grade_id: cell-e703cf046ffacac6, points: 0, grade: false, solution: true -->
```python
def answer_three():
    
    closeness_c = nx.closeness_centrality(G1)
    max_closeness_c = max(closeness_c, key=lambda x: closeness_c[x])
    
    return max_closeness_c

answer_three()
```

Output:
```
23
```

<!-- grade_id: cell-467293083292b0ee, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

### Question 4

Assume the restriction on the voucher’s travel distance is still removed, but now a competitor has developed a strategy to remove a person from the network in order to disrupt the distribution of your company’s voucher. You competitor plans to remove people who act as bridges in the network. Identify the best possible person to be removed by your competitor?

*This function should return an integer, the chosen node.*

<!-- grade_id: cell-0aa27c6d1caf2520, points: 0, grade: false, solution: true -->
```python
def answer_four():
    
    betweenness_c = nx.betweenness_centrality(G1)
    max_betweenness_c = max(betweenness_c, key=lambda x: betweenness_c[x])
    
    return max_betweenness_c

answer_four()
```

Output:
```
333
```

<!-- grade_id: cell-338a3ca88864385b, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()

```

```python
import matplotlib.pyplot as plt


#node_color = [G1.degree(v) for v in G1]
node_size = [nx.betweenness_centrality(G1)[v] for v in G1]

pos = nx.spring_layout(G1)
nx.draw_networkx(G1, pos, node_size=node_size, width=0.2, font_size=2, cmap=plt.cm.Blues)

```

Output:
```
File /opt/conda/lib/python3.9/site-packages/networkx/algorithms/centrality/betweenness.py:130, in betweenness_centrality(G, k, normalized, weight, endpoints, seed)
    127 for s in nodes:
    128     # single source shortest paths
    129     if weight is None:  # use BFS
--> 130         S, P, sigma, _ = _single_source_shortest_path_basic(G, s)
    131     else:  # use Dijkstra's algorithm
    132         S, P, sigma, _ = _single_source_dijkstra_path_basic(G, s, weight)

File /opt/conda/lib/python3.9/site-packages/networkx/algorithms/centrality/betweenness.py:277, in _single_source_shortest_path_basic(G, s)
    275         if D[w] == Dv + 1:  # this is a shortest path, count paths
    276             sigma[w] += sigmav
--> 277             P[w].append(v)  # predecessors
    278 return S, P, sigma, D

KeyboardInterrupt: 
<Figure size 1000x800 with 0 Axes>
```

## Part 2

`G2` is a directed network of political blogs, where nodes correspond to a blog and edges correspond to links between blogs. Use your knowledge of PageRank and HITS to answer Questions 5-9.

```python
G2 = nx.read_gml('assets/blogs.gml')
```

### Question 5

Apply the Scaled Page Rank Algorithm to this network. Find the Page Rank of node 'realclearpolitics.com' with damping value 0.85.

*This function should return a float.*

<!-- grade_id: cell-29a0e65bd30b8324, points: 0, grade: false, solution: true -->
```python
def answer_five():
    
    return nx.pagerank(G2, alpha=0.85)['realclearpolitics.com']
    
answer_five()
```

Output:
```
0.004636694781649098
```

<!-- grade_id: cell-5ade64a5771dcbce, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

### Question 6

Apply the Scaled Page Rank Algorithm to this network with damping value 0.85. Find the 5 nodes with highest Page Rank. 

*This function should return a list of the top 5 blogs in desending order of Page Rank.*

```python
nx.pagerank(G2, alpha=0.85)
```

Output:
```
{'tsrightdominion.blogspot.com': 0.0003478782612386359,
 'rightrainbow.com': 0.0004286700876475778,
 'gregpalast.com': 0.002149383652281393,
 'younglibs.com': 0.00021555645931858387,
 'blotts.org/polilog': 0.00018740248208077342,
 'marylandpolitics.blogspot.com': 0.00018740248208077342,
 'blogitics.com': 0.0002449721617548766,
 'thesakeofargument.com': 0.0002989949277473624,
 'joebrent.blogspot.com': 0.00018740248208077342,
 'thesiliconmind.blogspot.com': 0.00018740248208077342,
 'kippsblog.com': 0.00018740248208077342,
 'progressivetrail.org/blog.shtml': 0.00018740248208077342,
 'randomjottings.net': 0.00043947432143907157,
 'rightvoices.com': 0.00043791575641368514,
 '84rules.blog-city.com': 0.00018740248208077342,
 'blogs.salon.com/0002874': 0.00220422764231655,
 'alvintostig.typepad.com': 0.0003620007262167064,
 'notgeniuses.com': 0.0012257303648885753,
 'democratreport.blogspot.com': 0.0001974475727209362,
 'victorhanson.com': 0.0022248558929737496,
 'aldaron.modblog.com': 0.00018
```

<!-- grade_id: cell-dc31ae880c39420f, points: 0, grade: false, solution: true -->
```python
def answer_six():
    
    p_rank = nx.pagerank(G2, alpha=0.85)
    
    sorted_p_rank = sorted(p_rank.items(), key=lambda x: x[1], reverse=True)[0:5]
    
    top5 = list(zip(*sorted_p_rank))[0]
    
    return list(top5)

answer_six()
```

Output:
```
('dailykos.com',
 'atrios.blogspot.com',
 'instapundit.com',
 'blogsforbush.com',
 'talkingpointsmemo.com')
```

<!-- grade_id: cell-fa118135cb4998f4, points: 1, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == list, "You must return a list"

```

### Question 7

Apply the HITS Algorithm to the network to find the hub and authority scores of node 'realclearpolitics.com'. 

*Your result should return a tuple of floats `(hub_score, authority_score)`.*

<!-- grade_id: cell-bc1c0cf4f5f4f66a, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    
    hub = nx.hits(G2)[0]['realclearpolitics.com']
    authority = nx.hits(G2)[1]['realclearpolitics.com']
    
    return (hub, authority)

answer_seven()
```

Output:
```
(0.0003243556140278736, 0.003918957644934252)
```

<!-- grade_id: cell-43b3de064e549ef6, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()
assert type(ans_seven) == tuple, "You must return a tuple"

```

### Question 8 

Apply the HITS Algorithm to this network to find the 5 nodes with highest hub scores.

*This function should return a list of the top 5 blogs in desending order of hub scores.*

<!-- grade_id: cell-8f45f0320e6c645b, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    
    hub_sc = nx.hits(G2)[0]
    
    sorted_hub_sc = sorted(hub_sc, key=lambda x: hub_sc[x], reverse=True)
    
    return sorted_hub_sc[0:5]

answer_eight()
```

Output:
```
['politicalstrategy.org',
 'madkane.com/notable.html',
 'liberaloasis.com',
 'stagefour.typepad.com/commonprejudice',
 'bodyandsoul.typepad.com']
```

<!-- grade_id: cell-72499b780b38eb2c, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()
assert type(ans_eight) == list, "You must return a list"

```

### Question 9 

Apply the HITS Algorithm to this network to find the 5 nodes with highest authority scores.

*This function should return a list of the top 5 blogs in desending order of authority scores.*

<!-- grade_id: cell-2c2606d7aaaae35a, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    
    aut_sc = nx.hits(G2)[1]
    
    sorted_aut_sc = sorted(aut_sc.keys(), key=lambda x: aut_sc[x], reverse=True)
    
    return sorted_aut_sc[0:5]

answer_nine()
```

Output:
```
['dailykos.com',
 'talkingpointsmemo.com',
 'atrios.blogspot.com',
 'washingtonmonthly.com',
 'talkleft.com']
```

<!-- grade_id: cell-bbc73cedc13c80ca, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == list, "You must return a list"

```